<a href="https://colab.research.google.com/github/mahmud-nobe/Basic-Problems-Function/blob/master/Colab%20Testing.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
!apt-get install texlive texlive-xetex texlive-latex-extra pandoc
!pip install pypandoc
```

Output:
```
Reading package lists... Done
Building dependency tree       
Reading state information... Done
pandoc is already the newest version (1.19.2.4~dfsg-1build4).
pandoc set to manually installed.
The following package was automatically installed and is no longer required:
  libnvidia-common-430
Use 'apt autoremove' to remove it.
The following additional packages will be installed:
  fonts-droid-fallback fonts-lato fonts-lmodern fonts-noto-mono fonts-texgyre
  libcupsfilters1 libcupsimage2 libgs9 libgs9-common libijs-0.35 libjbig2dec0
  libkpathsea6 libpotrace0 libptexenc1 libruby2.5 libsynctex1 libtexlua52
  libtexluajit2 libzzip-0-13 lmodern poppler-data preview-latex-style rake
  ruby ruby-did-you-mean ruby-minitest ruby-net-telnet ruby-power-assert
  ruby-test-unit ruby2.5 rubygems-integration t1utils tex-common tex-gyre
  texlive-base texlive-binaries texlive-fonts-recommended texlive-latex-base
  texlive-latex-recommended texlive-pictures texlive-plain-generic tipa
Suggested packages:
```

```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Drive already mounted at /content/drive; to attempt to forcibly remount, call drive.mount("/content/drive", force_remount=True).
```

```python
%pwd
```

Output:
```
'/content/drive/My Drive/Colab Notebooks'
```

```python
%ls
```

Output:
```
[0m[01;34mdrive[0m/  [01;34msample_data[0m/
```

```python
%cd drive
```

Output:
```
/content/drive
```

```python
%ls
```

Output:
```
[0m[01;34m'My Drive'[0m/  [01;34m'Shared drives'[0m/
```

```python
%cd Colab Notebooks 
```

Output:
```
/content/drive/My Drive/Colab Notebooks
```

```python
%ls
```

Output:
```
'3-body problem'
'Assignment 2.ipynb'
'Copy of cs110_pcw_2_1.ipynb'
'Copy of Demo.ipynb'
'Copy of Mahmud: CS112 Final Project (1).ipynb'
'Copy of Mahmud: CS112 Final Project.ipynb'
'Copy of Welcome To Colaboratory'
 cs110_pcw_1_2.ipynb
'CS112 Final Project.ipynb'
'httr testing.ipynb'
 insertionsort_pseudo.png
'Matching Assignment 4..ipynb'
 Matching_Assignment_4_Large_Dateset.ipynb
'NS111 Hubble Law'
'Peacekeeping workout.ipynb'
 selectionsort_pseudo.png
 Untitled
 Untitled0.ipynb
```

```python
!cp 3-body problem.ipynb ./
```

Output:
```
cp: cannot stat '3-body': No such file or directory
cp: cannot stat 'problem.ipynb': No such file or directory
```

```python
!cp Colab Notebooks/Copy of cs110_pcw_2_1.ipynb ./
```

Output:
```
cp: cannot stat 'Colab': No such file or directory
cp: cannot stat 'Notebooks/Copy': No such file or directory
cp: cannot stat 'of': No such file or directory
cp: cannot stat 'cs110_pcw_2_1.ipynb': No such file or directory
```

```python
%pwd
```

Output:
```
'/content/drive/My Drive/Colab Notebooks'
```

```python
%cp Copy of cs110_pcw_2_1.ipynb
```

Output:
```
cp: target 'cs110_pcw_2_1.ipynb' is not a directory
```

```python
!jupyter nbconvert --to PDF "cs110_pcw_1_2.ipynb"
```

```python

```

```python

```

```python
%ls
```

```python

```

```python
!jupyter nbconvert --to PDF "cs110_assignment_1.ipynb"
```

```python

```

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---


# CS110 Spring 2020- Assignment 1
# Divide and Conquer Sorting Algorithms

This assignment focuses on the implementation of sorting algorithms and analyzing their performance both mathematically (using theoretical arguments on the asymptotic behavior of algorithms ) and experimentally (i.e., running experiments for different input arrays and plotting relevant performance results). 

Every CS110 assignment begins with a check-up on your class responsibilities and professional standing. If you have submitted make-up work, you will also receive formative feedback on it, accompanied by a grade reflecting on your ability to address one of the course LOs. Thus to complete the first part of this assignment, you will need to take a screenshot of your CS110 dashboard on Forum where the following is visible:

* your name.
* your absences for the course have been set to excused up to the last session from week 2 (inclusively) .

This will be evidence that you have submitted acceptable pre-class and make-up work for a CS110 session you may have missed. Please check the specific CS110 make-up and pre-class policies in the syllabus of the course.


**NOTES:**

1. Your assignment submission needs to include the following resources:

* A PDF file must be the first resource. This file must be generated from the template notebook where you have written all of the answers (check this [link](https://docs.google.com/document/d/1gRMol9Ebbvyu1mvEKzma92o_N7ZbNXsPIb1QdQV0TeE/edit#heading=h.10jakf1cwxpq) for instructions on how to do this). Make sure that the PDF displays properly (all text and code can be seen within the paper margins), and that your work is neat and clearly presented. 
* Your second resource must be the template notebook you have downloaded from the gist provided and where you included your answers. The name of this resource must include your own name. Submit this file directly following the directions in this picture:


<img src="https://github.com/mahmud-nobe/cs110_assignment_1/blob/master/upload.png?raw=1" width="800" height="200">

If you are unable to submit the jupyter notebook directly, please compress all the materials of your submission, and submit the zip file as your second resource instead.

2. Tasks (1)-(7) will be graded on the indicated LOs; please make sure to consult their descriptions and rubrics in the course syllabus. You will not be penalized for not attempting the optional challenge.

3. In this first assignment, you will receive formative feedback on your reference, application and justification of the HCs but no actual grades. Only in subsequent assignments will the HCs applications be graded. This is an opportunity for you to deepen your argumentation skills without generating a grade. 

As such, after completing the assignment, evaluate the application of the HCs you have identified prior to and while you were working on this assignment and footnote them (refer to [these guidelines](https://docs.google.com/document/d/1s7yOVOtMIaHQdKLeRmZbq1gRqwJKfezBsfru9Q6PcHw/edit) on how to incorporate HCs in your work). 

Here are some examples of weak applications of some of the relevant HCs:

* Example 1: “#algorithms: I wrote an implementation of Bubble sort”. 
    * This is an extremely superficial use of the HC in a course on Algorithms, and your reference will be graded accordingly. Instead, consider what constitutes an algorithm (see Cormen et al., sections 1.1 and 1.2). Once you have a good definition of an algorithm, think of how this notion helped you approach the implementation of the algorithm, analyze its complexity and understand why it’s important to write an optimal Python implementation of the algorithm.

* Example 2: “#dataviz: I plotted nice curves showing the execution time of bubble sort, or I plotted beautiful curves with different colors and labels.”

    * Again, these two examples are very superficial uses of the HC #dataviz. Instead, consider writing down how do the plots and figures helped you interpret, analyze and write concluding remarks from your experiments. Or write about any insight you included in your work that came from being able to visualize the curves.


* Example 3: “#professionalism: I wrote a nice paper/article that follows all the directions in this assignment.” 

    * By now, you should realize that this is a poor application of the HC #professionalism. Instead, comment on how you actively considered the HC while deciding on the format, length, and style for writing your report.

4. Your code will be tested for similarity using Turnitin, both to other students’ work and examples available online. As such, be sure to cite all references that you used in devising your solution. Any plagiarism attempts will be referred to the ASC, as per the course’s policy in the syllabus.


** Complete the following tasks which will be graded in the designated LOs and foregrounded HCs (please consult the LOs rubrics on the syllabus):**



## Question 1. [HCs #responsibility; appropriate LO]

Submit a screenshot of your CS110 dashboard with the information described above. You can do this by writing:
```
from IPython.display import Image
Image(filename=’your_screenshot.png’)
```

My Screenshot:

<img src="https://github.com/mahmud-nobe/cs110_assignment_1/blob/master/course%20stat.PNG?raw=1">

## Question 2. [#SortingAlgorithms, #PythonProgramming, #CodeReadability] 

Write a Python 3 implementation of the three-way merge sort algorithm discussed in class using the code skeleton below. You should also provide at least three test cases (possibly edge cases) that demonstrate the correctness of your code. Your output must be a sorted **Python list.**


<!-- grade_id: cell-47961f0da38f1381, points: 0, grade: false, solution: true -->
```python
def three_way_merge(arr):
    """Implements three-way marge sort
    
    Input:
    arr: a Python list OR numpy array (your code should work with both of these data types)
    
    Output: a sorted Python list"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-41b1d220414c9bd0, points: 1, grade: true, solution: false -->
```python
assert(three_way_merge([4,3,2,1]) == [1,2,3,4])
assert(three_way_merge(np.array([4,3,2,1])) == [1,2,3,4]) 
```

## Question 3.  [#SortingAlgorithms, #PythonProgramming, #CodeReadability, #ComputationalCritique] 

Implement a second version of a three-way merge sort that calls bubble sort when sublists are below a certain length (of your choice) rather than continuing the subdivision process. Justify on the basis of theoretical and potentially also experimental arguments what might be an appropriate threshold for the input array for applying bubble sort.


<!-- grade_id: cell-e1e72cc08eb6a5f0, points: 0, grade: false, solution: true -->
```python
def extended_three_way_merge(arr):
    """Implements the second version of a three-way merge sort
    
    Input:
    arr: a Python list OR numpy array (your code should work with both of these data types)
    
    Output: a sorted Python list"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8ab1cb1e1b4c76fb, points: 0, grade: true, solution: false -->
```python
assert(extended_three_way_merge([4,3,2,1]) == [1,2,3,4])
assert(extended_three_way_merge(np.array([4,3,2,1])) == [1,2,3,4]) 
```

## Question 4 [#SortingAlgorithms, #PythonProgramming, #CodeReadability] 

Bucket sort (or Bin sort) is an algorithm that takes as inputs an n-element array and the number of buckets, k, to be used during sorting.  Then, the algorithm distributes the elements of the input array into k-different buckets and proceeds to sort the individual buckets. Then, merges the sorted buckets to obtained the sorted array. Here is pseudocode for the BucketSort algorithm:

<img src="https://github.com/mahmud-nobe/cs110_assignment_1/blob/master/bucket.png?raw=1" width="800" height="200">

The BucketSort code above calls the function GetBucketNum  (see the pseudocode below) to distribute all the elements of array A into k-buckets.  Every element in the array is assigned a bucket number based on its value (positive or negative numbers). GetBucketNum returns the bucket number that corresponds to element A[i].   It takes as inputs the element of the array, A[i], the max and min elements in A, the size of the intervals in every bucket (e.g., if you have numbers with values between  0 and 100 numbers and 5 buckets, every bucket has an interval of size 20 = [100-0]/5).  Notice that in pseudocode, the indices of the arrays are from 1 to n. Thus, GetBucketNum consistently returns a number between 1 and n (make sure you account for this in your Python program).


<img src="https://github.com/mahmud-nobe/cs110_assignment_1/blob/master/getbucketnum.png?raw=1" width="800" height="200">

Write a Python 3 implementation of BucketSort that uses the bubble sort algorithm for sorting the individual buckets in line 10 of the algorithm.

<!-- grade_id: cell-4f283016e5365a91, points: 0, grade: false, solution: true -->
```python
def bucket_sort(arr, k):
    """Implements BucketSort
    
    Input:
    arr: a Python list OR numpy array (your code should work with both of these data types)
    k: int, number of buckets
    
    Output: a sorted Python list"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2a287ffea077a4e6, points: 1, grade: true, solution: false -->
```python
assert(bucket_sort([4,3,2,1]) == [1,2,3,4])
assert(bucket_sort(np.array([4,3,2,1])) == [1,2,3,4])
```

## Question 5 [#SortingAlgorithms, #PythonProgramming, #CodeReadability]

Implement a second version of the BucketSort algorithm. This time in line 10 of BucketSort apply the Bucket sort algorithm recursively until the size of the bucket is less than or equal to k, the base case (notice that fewer than k partitions will be inefficient).


<!-- grade_id: cell-d90434b2aa46d9aa, points: 0, grade: false, solution: true -->
```python
def extended_bucket_sort(arr, k):
    """Implements the second version of the BucketSort algorithm
    
    Input:
    arr: a Python list OR numpy array (your code should work with both of these data types)
    k: int, number of buckets
    
    Output: a sorted Python list"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a42f473f6bd53e24, points: 1, grade: true, solution: false -->
```python
assert(extended_bucket_sort([4,3,2,1]) == [1,2,3,4])
assert(extended_bucket_sort(np.array([4,3,2,1])) == [1,2,3,4])
```

## Question 6 [#ComplexityAnalysis, #ComputationalCritique] 

Analyze and compare the practical run times of regular merge sort (i.e., two-way merge sort), three-way merge sort, and the augmented merge sort from (3) by producing a plot that illustrates how every running time grows with input size. Make sure to:

1. define what each algorithm's complexity is
2. enumerate the explicit assumptions made to assess each run time of the algorithm's run time.
3. and compare your benchmarks with the theoretical result we have discussed in class.


<!-- grade_id: cell-7de481a7d24e508b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 7. [#ComplexityAnalysis, #ComputationalCritique] 

Analyze and compare the practical run times of regular merge sort (i.e., two-way merge sort), Bucket sort, and extended_bucket_sort by producing a plot that illustrates how the algorithms' runtimes depend on both the input size and the nature of the input (e.g., totally randomized inputs vs inversely sorted inputs). Make sure to:

1. define what each algorithm's complexity is in terms of the BigO notation
2. enumerate the explicit assumptions made to assess each algorithm's run time
3. analyze the running times on inputs of at least three different natures
4. and compare your benchmarks with the theoretical result we have discussed in class.


<!-- grade_id: cell-76dce7d4abf57157, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## [Optional challenge] Question 8 (#SortingAlgorithm and/or #ComputationalCritique) 

Implement k-way merge sort, where the user specifies k. Develop and run experiments to support a hypothesis about the “best” value of k.


<!-- grade_id: cell-30f70f3f9d3ec067, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```